# 01: Sampling

CBE 20258. Numerical and Statistical Analysis. Spring 2020.

&#169; University of Notre Dame.

## Lecture 15 Learning Objectives

After studying this notebook and your notes, completing the activities, asking questions in class, you should be able to:
* Give at least three examples of how an engineer would use statistics
* Compute descriptive statistics for data using Pandas (Python package)
* Explain what *statistical independence* means.
* Interpret the elements of a covariance matrix
* Explain why some data distributions have a very different median and mean

<div style="background-color: rgba(255,0,0,0.05) ; padding: 10px; border: 1px solid darkred;"> 
<b>Note</b>: The home activities are optional and will count as extra credit. They are due Monday, March 16 @ 10am via Vocareum. You already completed many of these home activities in "Class 3: Numpy, Matplotlib, and Pandas". We recommend revisiting these as practice during spring break.
</div>

## 15a. Samlping

**Further Reading**: ยง1.1 in Navidi (2015)

### Basic Vocabulary and an Example

Here is a central procedure that underpins the entire field of statistics:
1. Collect a sample of data.
2. Analyze the data to infer something about a population.

We will make this concrete with an example. Imagine we are researching nutrition and health of college students. We want to know the average caloric intake (calories eaten per day) of all entire undergraduate population at Notre Dame.

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: Brainstorm at least two challenges with measuring the caloric intake of the entire student population.
</div>

**Activity Answer**:

For the reasons you just listed and more, we decide it is infeasible to accurately record everything that everyone on campus eats. Instead we decide to randomly select a sample of students and have them record a food diary. Of course, we'll have to train them on how to make accurate measurements, compensate them for their time, and take additional steps at estimate how accuracy of their journal entries.

Using this example, let's introduce some vocaburaly:

| Definition | Discussion for Our Example |
| - | - |
| **Population**: The entire collection of objects or outcomes about which information is sought. | We want to know about the eating habits of the entire **ND undergraduate student body**. |
| **Sample**: A subset of the population, containing the objects or outcomes that are actually observed. | We ask a **small group of ND undergraduates** to keep a food journal. |
| **Simple Random Sample**: A sample that is randomly choosen, where each collection of population items is equally likely to make up the sample. | We decide to randomly select $n$ undergraduate students by **drawing names of our a well-mixed hat**. |
| **Convenience Sample**: A sample that is obtained is some convienent way, and not drawn from a well-defined random method. | We decide to **ask a few classes** on campus to keep food journals. |

A fundamental limitations of convience samples is that they often introduce systematic biases.

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: In two sentences, describe how asking only a few classes to keep food journals could lead us to misguided conclusions about the entire population of ND undergraduates. How would using a random sample help safegaurd against (some) of these biases?
</div>

**Activity Answer**:

Alright, so we randomly select 30 undergraduates and strap Go-Pros to their heads. We painstakingly analyze the videos to verify their food journals are 100% accurate. We then calculate the average caloric intake is 2023.2 kilocalories per day.

Is this the true average of the population? In other words, if we somehow measured the caloric intake of the entire ND undergraduate student population, should we expect to the average to be exactly 2023.2 kcal/day?

No. Why? Because our sample was randomly selected. Thus we expect to see **sampling variation**. In other words, if we collected another sample from 30 other randomly selected students, we expect a slightly different result. We will soon see how **probability theory** gives us a mathematical framework to quantify this variability.

### Conceptual versus Tangible Populations

There are two types of populations:

| Type of Population | Example |
| - | - |
| **Tangible Population**: A population that is finite. Often a tangigle population decreases by one after an object is sampled. | Population of students. |
| **Conceptual Population**: A population that is infinite. | An engineer measures the concentration of a mixture five times. The population is the set of all possible outcomes of these five measurements. |

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: With a partner, for each of the following scenarions, i) define the population and ii) state whether it is conceptual or tangible.
</div>


1. A chemical process is run 15 times, and the yield is measured 15 times.
2. In a clinical trial to test a new drung that is designed to lower cholesterol, 100 people with high cholesterol levels are recruited to try the new drug.

**Class Activity Answer**:

### Independence

The items in a sample are **independent** if knowing the values of some of the items does not help to predict the values of the others. Because tangigle populations are finite, samples taken without replacement are not strictly independent. For example, if we choose student "John Doe" to keep a food journal, we know that "John Doe" will not be chosen again (assuming everyone has a unique name and we are sampling without replacement). But we often treate simple random samples as independent if the sample contains less than 5% of the population.

## 15b. Working with Data using Panda

On Sakai, you'll find `Datasets-All-Examples-Navidi.zip`. This file, which I downloaded from [the publisher](http://highered.mheducation.com/sites/0073401331/student_view0/data_sets.html), contains all of the data for the examples and tables in our textbook. We'll use many of these datasets to illustrate key concepts in class.

Let's start with Tables 1.1 and 1.2 (pg. 21), which give **particulate matter (PM) emissions in g/gal** for 138 and 62 vehicles at low and high altitudes, respectively. Please take a moment to get out your textbook and glance at the tables. I'll wait.

Now let's load the data into Python. In this class, we will use `Pandas`, which is a super popular and easy to use package/library/module for organizing and manipulating data. Here is a highly recommended [10 minutes to pandas](https://pandas.pydata.org/pandas-docs/stable/getting_started/10min.html) getting started tutorial.

The code below reads in the first text file.

```python
import pandas as pd
low = pd.read_csv('table1-1.csv')
```

This creates a Pandas **dataframe**, which is stored in the variable `low`. We can easily print its contents to the screen.

```python
print(low)
```

```python
len(low)
```

The first row (vehicle) is numbered 0, which is perhaps not a surprise. We see there are 138 rows in the dataset, which matches what we expect: data for 138 vehicles at low altitude.

The output above is ugly. We can use the `.head()` and `.tail()` commands to look at only the first and last five entries.

```python
low.head()
```

```python
low.tail()
```

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
    <b>Optional Home Activity</b>: Load the high altitude data, which is stored in <tt>table1-2.csv</tt> into the Pandas dataframe <tt>high</tt>. Verify there are 62 rows. Use the head command to see the first few rows.
</div>

```python
### BEGIN SOLUTION
high = pd.read_csv('table1-2.csv')

print("Number of rows = ",len(high))

high.head()

### END SOLUTION
```

<!-- grade_id: 14b-high, points: 0.1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

assert len(high) == 62, "There should be 62 rows."

### END HIDDEN TESTS
```

## 14c. Summary Statistics

**Further Reading**: ยง1.2 in Navidi (2015)

### Two Types of Data: Numerical and Categorical

As engineers, you will encounter both **numerical** (quantitative) and **categorical** (qualitative) data.

Unfortunately, Example 1.8 from the textbook was not included in the data files on the publisher's website. But do not fear! We can recreate the table from a dictionary.

```python
# Store all of the data from Example 1.8 in a dictionary
# Notice the keys are the column names
my_dict = {"Specimen":[1, 2, 3, 4, 5], "Torque":[165, 237, 222, 255, 194],"Failure Location":['Weld','Beam','Beam','Beam','Weld']}

# Convert the dictionary into a Pandas dataframe
my_df = pd.DataFrame(my_dict)

# Print
print(my_df)

# Look at the first five entries
my_df.head()

# Profit???
```

Well that was easy.

In this example, we see that *Torque* is a numerical variable and *Failure Location* is a categorical variable.

We will now learn about statistics to summarize key characteristics of samples. Let's start by focusing on numerical variables.

### Sample Mean

The sample mean is the average of the sample.

Let $X_1$, $X_2$, ..., $X_n$ be the sample. The **sample mean** is

$$\bar{X} = \frac{1}{n} \sum_{i=1}^{n} X_i$$

Statisticians have many quirks, which make them the butts of many jokes. I would like to draw your attention to two:
1. A capital variable, such as $X_i$, is often a **random variable**. More on this next class.
2. Statisticians like to give variables decorations. Here $\bar{~}$ (bar) means average. Later in the semester we'll see that $\hat{~}$ (hat) means estimate.

It is really easy to calculate the sample mean with pandas:

```python
low.mean()
```

That output looks strange.

```python
type(low.mean())
```

Interesting. The command `low.mean()` does not return a floating point number. It instead returns a variable that is type `pandas.core.series.Series`. What if I just want the sample mean in a floating point number?

In pandas, we can access a column using its name. Recall here are the first five elements in the data set:

```python
low.head()
```

This data set only has one column. But if we want to extract that column, we write:

```python
low.PM
```

And to get the numeric mean, we simply write:

```python
low.PM.mean()
```

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
    <b>Optional Home Activity</b>: Calculate the sample mean for PM in the high altitude data set. Store your answer (a float) in the variable <tt>ans_14b_1</tt>.
</div>

```python
### BEGIN SOLUTION
ans_14b_1 = high.PM.mean()
print(ans_14b_1)
### END SOLUTION
```

<!-- grade_id: 14b-1, points: 0.2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import numpy as np
assert np.abs(ans_14b_1 - 6.596) < 1E-2, "Try again."
### END HIDDEN TESTS
```

### Sample Variance and Standard Deviation

While mean reflects the average of a sample, variance and standard deviation measure the spread.

Let $X_1$, ..., $X_n$ be a sample. The **sample variance** is

$$s^2 = \frac{1}{n-1} \sum_{i}^{n} (X_i - \bar{X})^2$$

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: If the measured quantitfy $X$ is velocity with units m/s, then what are the units of variance $s^2$?
</div>

**Activity Answer**:

Let's dig into the idea that variance measures the spread of the data set. Consider a synthetic (i.e., made up, artificial) data sets with two columns, A and B:

```python
my_data = pd.DataFrame({"A":[0, 0, 5, 10, 10], "B":[4, 4, 5, 6, 6]})
my_data.head()
```

Both columns have the same mean (average):

```python
my_data.mean()
```

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
        <b>Optional Home Activity</b>: Please take a minute to verify (on paper or in your head) that both data sets do in fact have a mean of 5.
</div>

Cool, pandas just calcutaed the means for both columns in our synthetic data set.

What about the variance of both columns? The variance formula,

$$s^2 = \frac{1}{n-1} \sum_{i}^{n} (X_i - \bar{X})^2$$

sums the squared different between each datum (a.k.a., data point) and the mean. Thus we expect a data set with a larger spread to have a larger variance.

```python
my_data.var()
```

And this is in fact what we see with the variance calculation. Column A, with range 0 to 10, has a much larger variance than column B, which has a range of 4 to 6.

Often we prefer to work with **sample standard derivation**, which is the square root of **sample variance**:

$$s = \sqrt{ \frac{1}{n-1} \sum_{i}^{n} (X_i - \bar{X})^2}$$

Notice that $s$ has the same units as $X$.

```python
my_data.std()
```

But why do the standard deviation and variance formulas divide by $n-1$ and not $n$? Short answer: often, we do not know the population standard deviation. (Recall, the sample was drawn from a population.) So we use $s$ as an estimate for the population standard deviation. This uses one degree of freedom, so we divide by $n-1$ instead of $n$. The $-1$ is because we want to estimate one parameter (population standard deviation). We will revist this idea a few times during the semester. A common exercise in a graduate level statistics course is to prove that dividing by $n-1$ makes $s^2$ an unbias estimate of population variance. Still curious? Check out this video: https://www.youtube.com/watch?v=D1hgiAla3KI

If we really wanted to, we could change the degree of freedom from the default 1 to any number:

```python
# divide by n - 1 in variance and standard derviation formulae. This is the default
print("variance\n",my_data.var(ddof=1),"\n")
print("standard deviation\n",my_data.std(ddof=1))
```

```python
# divide by n in variance and standard derviation formulae.
print("variance\n",my_data.var(ddof=0),"\n")
print("standard deviation\n",my_data.std(ddof=0))
```

```python
# divide by n - 2 in variance and standard derviation formulae.
print("variance\n",my_data.var(ddof=2),"\n")
print("standard deviation\n",my_data.std(ddof=2))
```

Like with the mean, we can easily extract a floating point number from pandas.

```python
my_data["A"].std()
```

```python
my_data.A.std()
```

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
    <b>Optional Home Activity</b>: Calculate the standard deviation (using $n-1$) for the particulate matter example. Store the results in the dictionary <tt>ans_14b_2</tt> with keys "low" and "high". Hint: You'll need to calculate the standard deviation as a float and then save the answers into the dictionary.
</div>

```python
### BEGIN SOLUTION
ans_14b_2 = {"low":low.PM.std(), "high":high.PM.std()}
print(ans_14b_2)
### END SOLUTION
```

<!-- grade_id: 14b-2, points: 0.2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import numpy as np
assert np.abs(ans_14b_2["low"] - 2.558) < 1E-2, "Standard deviation for low altitude is not correct."
assert np.abs(ans_14b_2["high"] - 4.519) < 1E-2, "Standard deviation for high altitude is not correct."
### END HIDDEN TESTS
```

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Optional Home Activity</b>: Show the following formulae are equivalent to the definitions for sample variance and sample standard deviation given above. <b>This is excellent practice for the next exam.</b>
</div>

$$s^2 = \frac{1}{n-1} \left( \sum_{i=1}^{n} X_i^2 - n \bar{X} \right)$$

$$s = \sqrt{\frac{1}{n-1} \left( \sum_{i=1}^{n} X_i^2 - n \bar{X} \right)}$$

### Sample Median

The sample median also measures the center of a data set. To compute the median, we order the data from smallest to largest and find the middle. For a data set with a even number of objects, we average the two middle dataum.

As you likely expect, pandas computes the median:

```python
my_data.median()
```

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
    <b>Optional Home Activity</b>: Calculate the median for the particulate matter example. Store the results in the dictionary <tt>ans_14b_3</tt> with keys "low" and "high". Hint: You'll need to calculate the median as a float and then save the answers into the dictionary.
</div>

```python
### BEGIN SOLUTION
ans_14b_3 = {"low":low.PM.median(), "high":high.PM.median()}
print(ans_14b_3)
### END SOLUTION
```

<!-- grade_id: 14b-3, points: 0.2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import numpy as np
assert np.abs(ans_14b_3["low"] - 3.18) < 1E-2, "Median for low altitude is not correct."
assert np.abs(ans_14b_3["high"] - 5.75) < 1E-2, "Median for high altitude is not correct."
### END HIDDEN TESTS
```

### Sample Mode and Range

The **sample mode** is the most common element in a sample.

Let's review samples A and B:

```python
my_data.head()
```

Recall, these two data sets have the same mean and median but different variances. Let's compute the mode:

```python
my_data.mode()
```

Interesting. In data set A, both there are two 0 elements and two 10 elements. Thus 0 and 10 are tied for the mode, and pandas returns two values. Likewise, 4 and 6 are tied for the mode in data set B.

Let's look at the particulate matter example together.

```python
low.PM.mode()
```

This suggests that both 1.11 and 1.63 are tied for the mode. Notice that `.mode()` returns a dictionary-like structure. We can access the first mode with the index (key) 0:

```python
low.PM.mode()[0]
```

And the second mode with index (key) 1:

```python
low.PM.mode()[1]
```

We can investigate further using the `.value_counts()` command:

```python
low.PM.value_counts()
```

This gives us the number of times each value repeats in the data set, sorted from most to least common.

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
    <b>Optional Home Activity</b>: Calculate the mode for the high elevation particulate matter example. Store the results in the float <tt>ans_14b_4</tt>. Hint: You'll need to either extract the float from the pandas output with key 0 or manually enter your answer.
</div>

```python
### BEGIN SOLUTION
ans_14b_4 = high.PM.mode()[0]
print(ans_14b_4)
### END SOLUTION
```

<!-- grade_id: 14b-4, points: 0.2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import numpy as np
assert np.abs(ans_14b_4 - 6.32) < 1E-2, "Mode for high altitude is not correct."
### END HIDDEN TESTS
```

The **sample range** is the different between the smallest and largest values in a sample. Pandas allows us to easily calculate these:

```python
# Inspect the data set
my_data.head()
```

```python
# identify smallest values
my_data.min()
```

```python
# identify largest values
my_data.max()
```

### Quartiles and Percentiles

The median is the 50%-ile of the data set. Half of the elements are above and half are below. But we can generalize this to any numeric value between 0% and 100%. You'll notice that the pandas documentation says **quantile** instead of **percentile**. These terms mean the same thing.

```python
# Calculate 50% quantile, a.k.a, 50%-ile
low.PM.quantile(0.5)
```

```python
# Verify this is the median
low.PM.median()
```

**Quartiles** divide the data set into quarters (four pieces). Quartiles are specific percentiles:

| Quartiles | Percentile / Quantile |
| - | - |
| 1st | 25% |
| 2nd | 50% |
| 3rd | 75% |

Let's look at the low altitude particulate matter example again:

```python
# Determine the number of entries
low.PM.count()
```

It contains 138 datum. So what happens if we want to compute the 52.1%-ile?

Pandas will give us a number:

```python
low.PM.quantile(0.521)
```

Notice the answer ends in ...000004. *What is going on?*

Because there are only 138 elements, the quantiles increment by 0.724...%

```python
1/138
```

Pandas is doing interpolation under the hood because there is not a datum exactly at the 52.1%-ile. The ...000004 ending is an artifact of inexact floating point arithmetic.

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
    <b>Optional Home Activity</b>: Calculate the 80%-ile for the high elevation particulate matter example. Store the results in the float <tt>ans_14b_5</tt>.
</div>

```python
### BEGIN SOLUTION
ans_14b_5 = high.PM.quantile(0.8)
print(ans_14b_5)
### END SOLUTION
```

<!-- grade_id: 14b-5, points: 0.1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import numpy as np
assert np.abs(ans_14b_5 - 8.822) < 1E-2, "Quantile not correct. Make sure you are computing for high altitude and 80%-ile."
### END HIDDEN TESTS
```

### Multivariate Data Example: Exam 1 Scores

The file `Exam-1-scores.csv` contains anonymized numeric scores from Exam 1 this semester. Let's open the file with Pandas.

```python
# Open the file
exam1_full = pd.read_csv("Exam_1_scores.csv")
```

```python
# Print entire dataframe
print(exam1_full)
```

We see 65 rows (students) and 20 columns. A perfect score was 60 points. Let's loop over the column names:

```python
for c in exam1_full.columns:
    print(c)
```

Let's make a new Pandas dataframe with only the problem totals.

```python
# create empty dictionary
new_data = {}
new_data['Total'] = exam1_full['Total Score']
new_data['P1'] = exam1_full['1-A'] + exam1_full['1-B']
new_data['P2'] = exam1_full['2-A-1'] + exam1_full['2-A-2'] + exam1_full['2-B']
new_data['P2'] += exam1_full['2-C-1'] + exam1_full['2-C-2'] + exam1_full['2-C-3']
new_data['P3'] = exam1_full['3-A'] + exam1_full['3-B-1'] + exam1_full['3-B-2'] + exam1_full['3-B-3']
new_data['P4'] = exam1_full['4-A-1'] + exam1_full['4-A-2'] + exam1_full['4-B']
new_data['P4'] += exam1_full['4-B'] + exam1_full['4-C-1'] + exam1_full['4-C-2']
new_data['P5'] = exam1_full['5-A'] + exam1_full['5-B']
```

```python
exam1 = pd.DataFrame(new_data)
```

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: Print the first and last five elements of the data set.
</div>

```python
# print top ("head") of data frame
exam1.head()
```

```python
# print bottom ("tail") of data frame

```

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: Compute the mean, median, mode, standard deviation, and quartiles using pandas.
</div>

```python
# mean

```

```python
# median

```

```python
# mode

```

```python
# standard deviation

```

```python
# 25%-ile
q25 = exam1.quantile(0.25)
print(q25)
```

```python
# 50%-ile

```

```python
# 75%-ile
```

Pandas offers a single function to compute these summary statistics.

```python
# all in one line
exam1.describe()
```

### Sample Covariance

Standard devation (and variance) measure the average squared distance from each element of a dataset and the mean. Standard deviation is for a single dimension.

Often, we want to know to the extent two variables are related. Let $X_1$, $X_2$, ..., $X_n$ and $Y_1$, $Y_2$, ..., $Y_n$ be the sample. Each pair ($X_i$, $Y_i$) corresponds to one experiment. For example, $X$ could be the effluent temperature and $Y$ could be the conversion for an adiabitic reactor. The experiment was repeated for $n$ trials.


The **sample covariance** if a generalization of variance from one to two dimensions:

$$s_{X,Y}^2 = \frac{1}{n-1} \sum_{i}^{n} (X_i - \bar{X})(Y_i - \bar{Y})$$

$\bar{X}$ and $\bar{Y}$ are the sample means for variables $X$ and $Y$.

If both $X_i$ and $Y_i$ tend to move together, i.e., both are either above ($Y_i > \bar{Y}$ when $X_{i} > \bar{X}$) or below ($Y_i < \bar{Y}$ when $X_{i} < \bar{X}$) their sample means, then $s_{X,Y}^2 >> 0$.

If they move in oppositive directions, i.e., $Y_i > \bar{Y}$ when $X_{i} < \bar{X}$ and $Y_i > \bar{Y}$ when $X_{i} < \bar{X}$, then $s_{X,Y}^2 << 0$.

If there is not a strong trend, then $s_{X,Y}^2 \approx 0$.

We can quickly compute covariance with Pandas.

```python
exam1.cov()
```

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: What are the units of sample covariance?
</div>

**Home Activity Answer:**

Often, it is more convinent to interpret **sample correlation**, which is sample covariance scaled by the standard deviation of each variable.

$$
r_{X,Y} = \frac{s_{X,Y}}{s_X \cdot s_Y}
$$

Recall, $s_{X,Y}$ is the covariance. $s_{X}$ and $s_{Y}$ are the standard deviations for variable $X$ and $Y$.

By construction, correlation is bounded between -1 and 1. During the remainded of the semester, we will see why the following rules hold:

| Correlation | Interpretation |
| - | - |
| $r_{X,Y} = -1.0$ | **Perfect negative correlation.** Samples $X_i$ and $Y_i$ lie exactly on a straight line with a negative slope. |
| $-1.0 < r_{X,Y} < 0$ | **Negative correlation.** If we fit a line to the data $X_i$ and $Y_i$, the slope would be negative.
| $r_{X,Y} = 0$ | **No correlation.** If we fit a line to the data $X_i$ and $Y_i$, the slope would be zero. |
| $0 < r_{X,Y} < 1$ | **Positive correlation.** If we fit a line to the data $X_i$ and $Y_i$, the slope would be positive.
| $r_{X,Y} = -1.0$ | **Perfect positive correlation.** Samples $X_i$ and $Y_i$ lie exactly on a straight line with a positive slope. |

(Sorry for the weird table wrapping.)

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: What are the units of sample correlation?
</div>

**Activity Answer:**

Let's look at the correlation for the exam1 data.

```python
exam1.corr()
```

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: Interpret correlation for the exam data. Why is the diagonal exactly 1.0. What would a negative correlation mean?
</div>

What happens if we first normalize the scores by the number of available points?

```python
# Make a copy of the DataFrame
exam1_norm = exam1.copy()

# Divide by the total number of available points
exam1_norm["P1"] = exam1_norm["P1"] / 8
exam1_norm["P2"] = exam1_norm["P2"] / 14
exam1_norm["P3"] = exam1_norm["P3"] / 10
exam1_norm["P4"] = exam1_norm["P4"] / 17
exam1_norm["P5"] = exam1_norm["P5"] / 14
exam1_norm["Total"] = exam1_norm["Total"] / 60
```

```python
# Compute covariance matrix
exam1_norm.cov()
```

```python
# Compute correlation matrix
exam1_norm.corr()
```

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: Why is the correlation matrix not changed by scaling?
</div>

### Summary Statistics for Categorical Variables

Categorical variables are qualititative. Recall our example from earlier:

```python
my_df.head()
```

Here failure location was a categorical variable. We cannot compute the median, standard deviation, or range for a categorical variable. Instead, we often compute **frequencies** by counting.

```python
my_df['Failure Location'].value_counts()
```

We can instead compute the **sample proportions** by normalizing (dividing by) the total number of samples.

```python
my_df['Failure Location'].value_counts(normalize=True)
```

### Summary Statistics and Population Parameters

Each sample statistic we have discussed (e.g., mean, median, standard deviation, covariance) has a counterpart for the population. We will call numerical summaries of samples **statistics** and numerical summaries of populations **parameters**. A central idea in data analysis is to use statistics to estimate/infer parameters. Often, we cannot directly measure a population. So instead we sample.

## 14d. Visualizing Data

**Further Reading**: ยง1.3 in Navidi (2015)

There are many ways to visualize data. This semester, we will focus on scatter plots, histograms, and boxplots. We will use matplotlib.

```python
import matplotlib.pyplot as plt
```

### Scatter Plots

We already learned how to make a scatter plot in matplotlib.

```python
plt.plot(exam1.P2,exam1.P3,marker='.',linestyle=" ")
plt.xlabel("Score for P2 [points]")
plt.ylabel("Score for P3 [points]")
plt.title("Exam 1 Scores")
plt.show()
```

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: Why do these data appear to line on a grid?
</div>

### Example: Old Faithful

To illustrate key concepts for plotting, we will explore Table 1.6 in Navidi (2015).

```python
old_faithful = pd.read_csv('table1-6.csv')
```

```python
old_faithful.head()
```

This data set has two columns:
1. The duration of dormant periods (in minutes).
2. Whether the previous erruption was *short* or *long*.

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
    <b>Home Activity</b>: Take a look at the nicely formatted table in the textbook (pg. 32).
</div>

### Histogram

Let's visualize the distribution of dormant periods using a histogram.

```python
plt.hist(old_faithful.Dormancy)
plt.xlabel("Dormancy (in minutes)")
plt.ylabel("Count")
plt.title("Old Faithful (Table 1.6)")
plt.show()
```

The horizontal axis of the histogram shows the continous variable of interest. The vertical axis shows the number of observations in each bin. By default, `matplotlib.pylot` uses equal size bins and automatically determines the number of bins. We can, of course, override these options.

https://matplotlib.org/3.1.0/api/_as_gen/matplotlib.pyplot.hist.html

https://matplotlib.org/3.1.0/gallery/statistics/hist.html

For example, let's use 11 bins that go from 40 to 95 minutes to match the textbook. Let's also make the bins 50% transparent (`alpha=0.5`). Finally, we will set `density=True` to normalize by the number of observations and set the color to purple.

```python
plt.hist(old_faithful.Dormancy, bins=11, range=(40, 95),alpha=0.5,density=True,color='purple')
plt.xlabel("Dormancy (in minutes)")
plt.ylabel("Density")
plt.title("Old Faithful (Table 1.6)")
plt.show()
```

This plot is **bimodal**. There is one mode between 50 and 55 minutes and a second mode between 80 and 85 minutes. Often, bimodal data are due to not accounting for an important scientific phenomena.

Let's split our data based on duration of the previous eruption. We will make two histograms.

```python
# find the elements where Eruption is 'Short'
i = old_faithful.Eruption == 'Short'

# Create the histogram for short
plt.hist(old_faithful.Dormancy[i], bins=11, range=(40, 95),alpha=0.5,density=True,color='red',label='Short')

# find the elements where Eruption is 'Long'
j = old_faithful.Eruption == 'Long'

# Create histogram for long
plt.hist(old_faithful.Dormancy[j], bins=11, range=(40, 95),alpha=0.5,density=True,color='blue',label='Short')

# Add labels
plt.xlabel("Dormancy (in minutes)")
plt.ylabel("Density")
plt.title("Old Faithful (Table 1.6)")
plt.show()
```

Ah ha, we can explain the two modes by accounting for the duration of the previous eruption.

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: Make a histogram of the low altitude emissions data set.
</div>

```python
# hint: you can access the data with 
# low.PM
```

### Boxplot

A **boxplots** visualizes the median, first and third quartiles, and any outliers. Here is the basic anatomy of a boxplot:

![boxplot](https://miro.medium.com/max/18000/1*2c21SkzJMf3frPXPAR_gZA.png)

The **interquartile range** is the difference between the third and first quartiles. Recall, between the first and third quartiles are the middle 50% of the data. Any observations 1.5 times the IQR beyond the first and third quartiles are considers **extreme outliers** and marked with a $*$ or circle on the boxplot.

Let's make a boxplot of the Old Faithful data.

```python
# Notice that both the data and labels are in lists,
# and the lists have the same length
plt.boxplot([old_faithful.Dormancy], labels=['All Data'])
plt.ylabel("Dormancy (in minutes)")
plt.title("Old Faithful (Table 1.6)")
plt.show()
```

The organge line is the median. There are no outliers shown on this boxplot.

Now let's make a comparative boxplot.

```python
# find the elements where Eruption is 'Short'
i = old_faithful.Eruption == 'Short'

# find the elements where Eruption is 'Long'
j = old_faithful.Eruption == 'Long'

# Create two side-by-side boxplots
plt.boxplot([old_faithful.Dormancy[i], old_faithful.Dormancy[j]], labels=['Short','Long'])
plt.ylabel("Dormancy (in minutes)")
plt.title("Old Faithful (Table 1.6)")
plt.show()
```

Now we see some outliers.

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Optional Home Activity</b>: After class, test your skills by making a comparative boxplot for the low and high altitude emssions data set. Compare your work to Figure 1.15 in Navidi.
</div>

### Visualization Practice: Exam Scores

```python
# Cumulative distribution function
n = np.arange(0,len(exam1_norm))/len(exam1_norm)
plt.plot(exam1_norm['Total'].sort_values()*100,n,color="black",label="Total",linewidth=2)
plt.plot(exam1_norm['P1'].sort_values()*100,n,color="red",label="P1",linewidth=2,linestyle="--")
plt.plot(exam1_norm['P2'].sort_values()*100,n,color="blue",label="P2",linewidth=2,linestyle="--")
plt.plot(exam1_norm['P3'].sort_values()*100,n,color="green",label="P3",linewidth=2,linestyle="--")
plt.plot(exam1_norm['P4'].sort_values()*100,n,color="purple",label="P4",linewidth=2,linestyle="--")
plt.plot(exam1_norm['P5'].sort_values()*100,n,color="orange",label="P5",linewidth=2,linestyle="--")
plt.xlabel("Score [%]")
plt.ylabel("Fraction of Class")
plt.legend()
plt.show()
```

```python
# Histogram

fig, axes = plt.subplots(nrows=6, ncols=1, figsize=(10,20))
ax0, ax1, ax2, ax3, ax4, ax5 = axes.flatten()


ax0.hist(exam1_norm['Total'].sort_values()*100,color="black",label="Total")
ax0.set_title('Total')

ax1.hist(exam1_norm['P1'].sort_values()*100,color="red",label="P1")
ax1.set_title('Problem 1')

ax2.hist(exam1_norm['P2'].sort_values()*100,color="blue",label="P2")
ax2.set_title("Problem 2")

ax3.hist(exam1_norm['P3'].sort_values()*100,color="green",label="P3")
ax3.set_title('Problem 3')

ax4.hist(exam1_norm['P4'].sort_values()*100,color="purple",label="P4")
ax4.set_title('Problem 4')

ax5.hist(exam1_norm['P5'].sort_values()*100,color="orange",label="P5")
ax5.set_title('Problem 5')

plt.xlabel("Score [%]")
plt.ylabel("Fraction of Class")
plt.tight_layout()
```

```python
### Boxplot
plt.boxplot(exam1_norm.transpose(), labels=exam1_norm.columns)
plt.ylabel("Score [%]")
plt.show()
```