# Python Basics III: Lists, Dictionaries, and Enumeration

**Reference**: Chapters 2 and 5 of *Computational Nuclear Engineering and Radiological Science Using Python*, R. McClarren (2018) 

## Learning Objectives

After studying this notebook, completing the activities, and asking questions in class, you should be able to:
* Create lists with numeric and non-numeric items and add to a previously created list.
* Enumerate a list using various syntax.
* Create dictionaries with numeric and non-numeric key-value pairs and add to a previously created dictionary.
* Enumerate a dictionary using various syntax.

## Lists

You can also do other fun stuff with for loops.  For instance, you can have the control variable take on non-numeric things:

```python
#silly hat code
hats = ["fedora","trilby","porkpie","tam o'shanter",
        "Phrygian cap","Beefeaters' hat","sombrero"]
days = ["Monday","Tuesday","Wednesday","Thursday",
        "Friday","Saturday","Sunday"]
count = 0
for today in hats:
    print("It is",days[count],"and I will wear a",today)
    count += 1
```

Output:
```
It is Monday and I will wear a fedora
It is Tuesday and I will wear a trilby
It is Wednesday and I will wear a porkpie
It is Thursday and I will wear a tam o'shanter
It is Friday and I will wear a Phrygian cap
It is Saturday and I will wear a Beefeaters' hat
It is Sunday and I will wear a sombrero
```

Notice what else we did here: we defined a list called `days` that contained strings for the names of the days of the week.  Inside our `for` loop we had a numeric variable that kept track of what number the day of the week was (0 for Monday in this case).  Then when we access days[count] it returns the string in position count.

We can go one step beyond and plan our wardrobe a month in advance using random numbers. The code below generates a random integer between 0 and 6:

```python
import random
my_random_number = round(random.uniform(0,6))
print(my_random_number)
```

Output:
```
1
```

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Home Activity</b>: Modify the code below to loop over 30 days. For each day, randomly select which hat to wear.
</div>

```python
#silly hat code
import random
hats = ["fedora","trilby","porkpie","tam o'shanter",
        "Phrygian cap","Beefeaters' hat","sombrero"]
days = ["Monday","Tuesday","Wednesday","Thursday",
        "Friday","Saturday","Sunday"]

# loop over 30 days
### BEGIN SOLUTION
for count in range(30):
### END SOLUTION

    # randomly select a hat to wear
    ### BEGIN SOLUTION
    hat_choice = round(random.uniform(0,6))
    ### END SOLUTION
    
    # print which day of the week it is and the hat choice
    ### BEGIN SOLUTION
    print("It is",days[count % 7],
          "and I will wear a",hats[hat_choice])
    ### END SOLUTION
```

Output:
```
It is Monday and I will wear a fedora
It is Tuesday and I will wear a Phrygian cap
It is Wednesday and I will wear a Phrygian cap
It is Thursday and I will wear a Phrygian cap
It is Friday and I will wear a tam o'shanter
It is Saturday and I will wear a porkpie
It is Sunday and I will wear a Phrygian cap
It is Monday and I will wear a Phrygian cap
It is Tuesday and I will wear a trilby
It is Wednesday and I will wear a sombrero
It is Thursday and I will wear a fedora
It is Friday and I will wear a trilby
It is Saturday and I will wear a sombrero
It is Sunday and I will wear a sombrero
It is Monday and I will wear a trilby
It is Tuesday and I will wear a porkpie
It is Wednesday and I will wear a porkpie
It is Thursday and I will wear a Beefeaters' hat
It is Friday and I will wear a tam o'shanter
It is Saturday and I will wear a trilby
It is Sunday and I will wear a Phrygian cap
It is Monday and I will wear a sombrero
It is Tuesday and I will wear a sombrero
It is Wednesday and I will we
```

**Lists** are like vectors in MATLAB, except, as you can see above, we can store more than just numbers in them. We will talk about using a package `numpy` to represent vectors and matrices in a future notebook.

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Home Activity</b>: Fix the three examples below to count from 1 to 10.
</div>

```python
## Attempt 1
N = [1, 2, 3, 4, 5, 6, 7, 11]
for i in range(len(N)):
    print(i+1)
```

Output:
```
1
2
3
4
5
6
7
8
```

```python
### BEGIN SOLUTION
# This is one of many solutions
N = [1, 1, 1, 1, 1, 1, 1, 1, 1, 1]
for i in range(len(N)):
    print(i+1)
### END SOLUTION
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

```python
## Attempt 2
N = [1, 2, 3, 4, 5, 6, 7, 11]
for i in range(len(N)):
    print(N[i])
```

Output:
```
1
2
3
4
5
6
7
11
```

```python
### BEGIN SOLUTION
N = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
for i in range(len(N)):
    print(N[i])
### END SOLUTION
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

```python
## Attempt 3
for i in range(0,10):
    print(i)
```

Output:
```
0
1
2
3
4
5
6
7
8
9
```

```python
### BEGIN SOLUTION
for i in range(1,11):
    print(i)
### END SOLUTION
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

## Dictionaries

So far, we have used lists to store a sequence of numbers or more generally Python objects. Lists (and NumPy arrays - stay tuned) access elements by position and have an inherent ordering.

A dictionary is a set of **key** : **value** pairs. You can think of a dictionary as a fancy list that instead of accessing via the position, you access using a key.

```python
#simple dictionary
days_of_week = {"M":"Monday", "T":"Tuesday",
                "W":"Wednesday", "R":"Thursday",
                "F":"Friday", "S":"Saturday",
                "U":"Sunday"}

print("Key M gives", days_of_week["M"])
print("Key R gives", days_of_week["R"])
```

Output:
```
Key M gives Monday
Key R gives Thursday
```

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Home Activity</b>: Store the value for key <tt>F</tt> in the variable <tt>answer</tt>.
</div>

```python
### BEGIN SOLUTION
answer = days_of_week["F"]
### END SOLUTION
```

<!-- grade_id: dictionary, points: 0.3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert answer == "Friday", "Your answer is not correct."
### END HIDDEN TESTS
```

Calling `days_of_week.keys()` returns a special list of the keys of the dictionary.

```python
print(days_of_week.keys())
```

Output:
```
dict_keys(['M', 'T', 'W', 'R', 'F', 'S', 'U'])
```

We can also check if `G` is a valid key for the dictionary.

```python
print("G" in days_of_week.keys()) #is G a key in days_of_week?
```

Output:
```
False
```

**Main idea:** With dictionaries, we access elements using the **key**. In contrast, we access elements of strings, lists, and Numpy arrays using the position.

### Example: Shopping List

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Home Activity</b>: Create a dictionary to store the following shopping list. The keys should be the item names and the elements should be the quantity of the items.
</div>

* 1.5 bananas (pounds)
* 3 cans of soup
* 6 apples
* 2.3 pork tenderloin (pounds)

```python
### BEGIN SOLUTION
shopping_list = {"bananas (pounds)":1.5, "cans of soup":3, "apples":6, "pork tenderloin (pounds)":2.3}
print(shopping_list)
### END SOLUTION
```

Output:
```
{'bananas (pounds)': 1.5, 'cans of soup': 3, 'apples': 6, 'pork tenderloin (pounds)': 2.3}
```

<!-- grade_id: 4a-i, points: 0.4, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert shopping_list["bananas (pounds)"] == 1.5
assert shopping_list["cans of soup"] == 3
assert shopping_list["apples"] == 6
assert shopping_list["pork tenderloin (pounds)"] == 2.3
### END HIDDEN TESTS
```

### Another Example: Chemical Symbols and Element Names

Below is code that parses ``ChemicalSymbols.csv`` and stores the data in dictionary.

```python
import csv
element_dict = {} #create a blank dictionary

# this block will only execute if the file opens
with open('../data/ChemicalSymbols.csv') as csvfile: 
    chemreader = csv.reader(csvfile)
    for row in chemreader: # have for loop that loops over each line
        element_dict[row[0]] = row[1] # add a key:value pair

# Ask user to enter a chemical symbol        
key = input("Enter a valid chemical symbol: ")
if key in element_dict:
    print(key,"is",element_dict[key])
else:
    print("Not a valid element")
```

Output:
```
Enter a valid chemical symbol:  K
K is Potassium
```

<div style="background-color: rgba(0,255,0,0.05) ; padding: 10px; border: 1px solid darkgreen;"> 
<b>Home Activity</b>: Run the code above twice. Try first with a valid chemical symbol. Then try with a chemical symbol not on the periodic table of elements.
</div>

We can also print out the dictionary.

```python
print(element_dict)
```

Output:
```
{'Ac': 'Actinium', 'Ag': 'Silver', 'Al': 'Aluminium (Aluminum)', 'Am': 'Americium', 'Ar': 'Argon', 'As': 'Arsenic', 'At': 'Astatine', 'Au': 'Gold', 'B': 'Boron', 'Ba': 'Barium', 'Be': 'Beryllium', 'Bh': 'Bohrium', 'Bi': 'Bismuth', 'Bk': 'Berkelium', 'Br': 'Bromine', 'C': 'Carbon', 'Ca': 'Calcium', 'Cd': 'Cadmium', 'Ce': 'Cerium', 'Cf': 'Californium', 'Cl': 'Chlorine', 'Cm': 'Curium', 'Cn': 'Copernicium', 'Co': 'Cobalt', 'Cr': 'Chromium', 'Cs': 'Caesium (Cesium)', 'Cu': 'Copper', 'Db': 'Dubnium', 'Ds': 'Darmstadtium', 'Dy': 'Dysprosium', 'Er': 'Erbium', 'Es': 'Einsteinium', 'Eu': 'Europium', 'F': 'Fluorine', 'Fe': 'Iron', 'Fl': 'Flerovium', 'Fm': 'Fermium', 'Fr': 'Francium', 'Ga': 'Gallium', 'Gd': 'Gadolinium', 'Ge': 'Germanium', 'H': 'Hydrogen', 'He': 'Helium', 'Hf': 'Hafnium', 'Hg': 'Mercury', 'Ho': 'Holmium', 'Hs': 'Hassium', 'I': 'Iodine', 'In': 'Indium', 'Ir': 'Iridium', 'K': 'Potassium', 'Kr': 'Krypton', 'La': 'Lanthanum', 'Li': 'Lithium', 'Lr': 'Lawrencium', 'Lu': 'Lutetium', 'Lv
```

```python
element_dict
```

Output:
```
{'Ac': 'Actinium',
 'Ag': 'Silver',
 'Al': 'Aluminium (Aluminum)',
 'Am': 'Americium',
 'Ar': 'Argon',
 'As': 'Arsenic',
 'At': 'Astatine',
 'Au': 'Gold',
 'B': 'Boron',
 'Ba': 'Barium',
 'Be': 'Beryllium',
 'Bh': 'Bohrium',
 'Bi': 'Bismuth',
 'Bk': 'Berkelium',
 'Br': 'Bromine',
 'C': 'Carbon',
 'Ca': 'Calcium',
 'Cd': 'Cadmium',
 'Ce': 'Cerium',
 'Cf': 'Californium',
 'Cl': 'Chlorine',
 'Cm': 'Curium',
 'Cn': 'Copernicium',
 'Co': 'Cobalt',
 'Cr': 'Chromium',
 'Cs': 'Caesium (Cesium)',
 'Cu': 'Copper',
 'Db': 'Dubnium',
 'Ds': 'Darmstadtium',
 'Dy': 'Dysprosium',
 'Er': 'Erbium',
 'Es': 'Einsteinium',
 'Eu': 'Europium',
 'F': 'Fluorine',
 'Fe': 'Iron',
 'Fl': 'Flerovium',
 'Fm': 'Fermium',
 'Fr': 'Francium',
 'Ga': 'Gallium',
 'Gd': 'Gadolinium',
 'Ge': 'Germanium',
 'H': 'Hydrogen',
 'He': 'Helium',
 'Hf': 'Hafnium',
 'Hg': 'Mercury',
 'Ho': 'Holmium',
 'Hs': 'Hassium',
 'I': 'Iodine',
 'In': 'Indium',
 'Ir': 'Iridium',
 'K': 'Potassium',
 'Kr': 'Krypton',
 'La': 'Lanthanum',
 'Li':
```

### Dictionary of Dictionaries

We can make dictionaries even more powerful if we make a dictionary of dictionaries.

**Top Level**. Day of the week.
* Key: symbol "M", "T", etc.
* Value: Another dictionary!

**Second Level**. Data for each day.
* Keys: "name", "weekday", "weekend"
* Values: string, boolean (True or False), boolean

```python
#simple dictionary of dictionaries
days_of_week = {"M":{"name":"Monday","weekday":True,"weekend":False},
                "T":{"name":"Tuesday","weekday":True,"weekend":False},
                "W":{"name":"Wednesday","weekday":True,"weekend":False},
                "R":{"name":"Thursday","weekday":True,"weekend":False},
                "F":{"name":"Friday","weekday":True,"weekend":False},
                "S":{"name":"Saturday","weekday":False,"weekend":True},
                "U":{"name":"Sunday","weekday":False,"weekend":True}}

print("The days that are weekdays are")
for day in days_of_week: #for loop over dictionary, loops over keys
    if days_of_week[day]["weekday"] == True:
        print(days_of_week[day]["name"],"is a weekday.")

print("The days that are weekend days are")        
for day in days_of_week: #for loop over dictionary, loops over keys
    if days_of_week[day]["weekend"] == True:
        print(days_of_week[day]["name"],"is a weekend, whoop.")
```

Output:
```
The days that are weekdays are
Monday is a weekday.
Tuesday is a weekday.
Wednesday is a weekday.
Thursday is a weekday.
Friday is a weekday.
The days that are weekend days are
Saturday is a weekend, whoop.
Sunday is a weekend, whoop.
```

Notice that when you loop over a dictionary in a for loop, the loop variable will get each of the keys of the dictionary.

The dictionary has been modified to include the day name in Spanish.

```python
#more complicated dictionary of dictionaries
days_of_week2 = {"M":{"name":"Monday","weekday":True,"Spanish":"Lunes"},
                "T":{"name":"Tuesday","weekday":True,"Spanish":"Martes"},
                "W":{"name":"Wednesday","weekday":True,"Spanish":"Miércoles"},
                "R":{"name":"Thursday","weekday":True,"Spanish":"Jueves"},
                "F":{"name":"Friday","weekday":True,"Spanish":"Viernes"},
                "S":{"name":"Saturday","weekday":False,"Spanish":"Sábado"},
                "U":{"name":"Sunday","weekday":False,"Spanish":"Domingo"}}
```

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: With a partner, write pseudocode to loop over the days of the week. Print off the day name in Spanish and whether it is a weekend or weekday.
</div>

```python

```

<div style="background-color: rgba(0,0,255,0.05) ; padding: 10px; border: 1px solid darkblue;"> 
<b>Class Activity</b>: Implement your pseudocode in Python.
</div>

```python

```

## Enumerate

### Lists

Let's see some syntax for growing lists.

```python
# create a list with nothing

students = []

# add a student
students.append("Joe Smith")

# add another student
students.append("Jane Doe")
```

```python
print(students)
```

Output:
```
['Joe Smith', 'Jane Doe']
```

How to loop over the list? Let's review.

```python
# Approach 1
for s in students:
    print(s)
```

Output:
```
Joe Smith
Jane Doe
```

```python
# Approach 2
for i in range(len(students)):
    print(i,": ",students[i])
```

Output:
```
0 :  Joe Smith
1 :  Jane Doe
```

Can we simply the syntax? It is really convienent to access the index with `i` and the student with `s`.

```python
for i, s in enumerate(students):
    print(i,": ",s)
```

Output:
```
0 :  Joe Smith
1 :  Jane Doe
```

### Dictionaries

How about adding key-value pairs to a dictionary? Let's take a look.

```python
# create a dictionary with nothing

pantry = {}

# add a pantry item and quantity
pantry["pasta"] = 3

# add another pantry item and quantity
pantry["green beans"] = 4
```

```python
print(pantry)
```

Output:
```
{'pasta': 3, 'green beans': 4}
```

If the key already exists in the dictionary, it will update with the newly assigned value rather than add another pair.

```python
pantry["pasta"] = 2
print(pantry)
```

Output:
```
{'pasta': 2, 'green beans': 4}
```

How do we loop over the dictionary? Here it is again.

```python
# to print out just the keys
for k in pantry.keys():
    print(k)
```

Output:
```
pasta
green beans
```

```python
# to print out just the values
for v in pantry.values():
    print(v)
```

Output:
```
2
4
```

```python
# to print out the key-value pairs together
for food in pantry:
    print(food, ":", pantry[food])
```

Output:
```
pasta : 2
green beans : 4
```